---
title: "A Minimal `{submitr}` Example"
output: learnr::tutorial
runtime: shiny_prerendered
tutorial:
  id: "minimal-example"
  version: 0.5
---

```{r setup, include=FALSE}
library(learnr)
library(submitr)
library(basket)
library(etude)
library(gradethis)
knitr::opts_chunk$set(echo = FALSE)
learnr::tutorial_options(
  exercise.timelimit = 60, 
  exercise.checker = #gradethis::grade_learnr) 
  submitr::null_code_checker)
```

```{r}
submitr::login_controls()
```

```{r context = "server", echo = FALSE}
options(tutorial.storage = "none")
vfun <- make_basic_validator(NULL, "hello") #basket::check_valid
storage_actions <- record_local("./minimal_submissions.csv")
submitr::shiny_logic(input, output, session, vfun,
                     storage_actions)
```

<!--End of submitr setup-->



## Fixed-response questions

```{r planets}
learnr::question("Which planet do we live on?",
         answer("Mars", correct=FALSE),
         answer("Earth", correct = TRUE,  message = "Good job Earthling!"),
         answer("Saturn",  correct = FALSE),
         allow_retry = TRUE,
         random_answer_order = TRUE
)
```

## Essay


```{r document_d52f3dd-Essay1, echo=FALSE}
etude::essay_response(
  prompt = "Write your thoughts here"
)
```

## Video

![Just a link](https://youtu.be/zNzZ1PfUDNk)


## Graded Code

```{r code_sandbox, exercise=TRUE, exercise.cap="Graded"}

```

```{r code_sandbox-check}
gradethis::grade_result(
  fail_if( ~ .result != 3, "Result should have been 3."),
  pass_if(~ TRUE, "Good for you!"),
  glue_correct = "{.message}"
)
```

## Ungraded Code

```{r ungraded_code, exercise = TRUE, exercise.cap = "Sandbox"}

```

## Multiple selections in multiple choice

```{r multiple-selections}
learnr::question("Which of these are countries?",
         answer("California", correct=FALSE),
         answer("France", correct = TRUE,  message = "Good job Earthling!"),
         answer("China",  correct = TRUE),
         random_answer_order = FALSE,
         options = list(foobar = "Flag")
)
```



